<a href="https://colab.research.google.com/github/Dilshodbek23/Data-Science-and-Artificial-Intelligence/blob/main/task(2_1_2)_Assignment_on_the_basics_of_indexing_and_cutting_arrays_ipynb%22.ipynb" target="_parent"><img src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open In Colab"/></a>

---
# Data analysis
2.1.1. Processing N-Dimensional Arrays

Task-2

```python
# calling a NumPy library
import numpy as np
```

### Task (1): Create a 2D Array (1 point)
Requerements: create an array of size (3, 3) with elements from 0 to 9 (9 itself is not included in the array element) and step 1.

Help: use `arange` to makes the job easier


<!-- grade_id: 3x3array, points: 0, grade: false, solution: true -->
```python
def nd_array():
    """Array of size 0 to 9 (9 is not an array element) and step equal to 1 (3, 3) """
    # YOUR CODE HERE
    arr = np.arange(9).reshape(3,3)
    return arr
    raise NotImplementedError()
```

Your function is tested by the number of elements in the array, the shape of the array `shape`, and whether the elements are multiple arrays in the range 0 ~ 10. There are also a number of secret tests (our advice is not to try to cheat).

<!-- grade_id: correct_3x3array, points: 1, grade: true, solution: false -->
```python
"""小heck the function"""
arr = nd_array() # massiv
assert arr.size == 9  # elementlar soni
assert type(arr) == np.ndarray # #numpy array 
```

###  Task (2): Cut elements from 2D arrays using an index (3 points)
Task: Create a function to cut out a portion of the array created above with elements 6 and 7.

Help: Find indexes 6 and 7 and shorten them using these indexes. Result: should be in the form of an array ([6, 7]).

<!-- grade_id: 2dsliced_array, points: 0, grade: false, solution: true -->
```python
def sliced_array_2d():
    """Function to trim part of an array with elements 6 and 7 """
    # YOUR CODE HERE
    sliced_arr = arr[2,:2]
    return sliced_arr
    raise NotImplementedError()
```

Your function is checked by the number of elements in the array, the shape of the array `shape`, and whether the elements are a numpy array. There are also a number of secret tests (our advice is not to try to cheat).

<!-- grade_id: correct_2dsliced_array, points: 3, grade: true, solution: false -->
```python
"""小heck the function"""
sl_arr_2d = sliced_array_2d() # truncated array
assert sl_arr_2d.size == 2  # number of elements
assert type(sl_arr_2d) == np.ndarray # #numpy array 
```

###  Task (3): Cut elements from a 3D array using the index (4 points)
Task: Create a 3D array from the list below, and cut out elements 12, 13, 15, and 16.

                  [[[ 0,  1,  2],
                    [ 3,  4,  5],
                    [ 6,  7,  8]],

                   [[ 9, 10, 11],
                    [12, 13, 14],
                    [15, 16, 17]],

                   [[18, 19, 20],
                    [21, 22, 23],
                    [24, 25, 26]]]

Help: First bring the array to a two-dimensional array, then perform the cut.

The result should look like this: 

             array([[12, 13],
                    [15, 16]])

Your function is checked by the number of elements in the array, the shape of the array `shape`, and whether the elements are numpy arrays. There are also a number of secret tests (our advice is not to try to cheat).

<!-- grade_id: 3dsliced_arr, points: 0, grade: false, solution: true -->
```python
def sliced_array_3d():
    """The function of creating a 3D array from the list above, as well as cutting elements 12, 13, 15, and 16"""
    # YOUR CODE HERE
    arr = np.arange(27).reshape(3,3,3)
    arr_2 = arr[1][1:,:2]
    return arr_2
    raise NotImplementedError()
```

<!-- grade_id: correct_3dsliced_arr, points: 4, grade: true, solution: false -->
```python
"""小heck the function"""
sl_arr_3d = sliced_array_3d() # truncated array
assert sl_arr_3d.size == 4  # number of elements
assert type(sl_arr_3d) == np.ndarray # #numpy array 
```

###  Task (4): Cut elements from 2-dimensional arrays using the Boolean index (2 points)
Case: We have ** listed people and their ** information ** in the form
        
        ['Hasan', 'Husan', 'Javohir', 'Elyor', 'Hasan', 'Javohir', 'Elyor']

       
      [[5, 6, 1, 1],
       [9, 1, 1, 1],
       [7, 7, 4, 2],
       [1, 5, 1, 9],
       [9, 9, 4, 5],
       [7, 5, 9, 6],
       [5, 3, 7, 4]]
 


From this data, create a function to cut the data for the names ** Javohir ** and ** Elyor ** using Boolean indexing.
** Help: ** Refer to the Boolean Indexing Video Tutorial. ** Result: ** The result is in the form of a dimensional array (4, 4), and the first two rows are as follows

    array([[7, 7, 4, 2],
           [1, 5, 1, 9],
           ...
           ...        ]])

<!-- grade_id: boolean_idx, points: 0, grade: false, solution: true -->
```python
def boolean_slicing():
    """Cutting data using boolean indexing """
    # YOUR CODE HERE
    name = np.array(['Hasan', 'Husan', 'Javohir', 'Elyor', 'Hasan', 'Javohir', 'Elyor'])
    data = np.array([[5, 6, 1, 1],
   [9, 1, 1, 1],
   [7, 7, 4, 2],
   [1, 5, 1, 9],
   [9, 9, 4, 5],
   [7, 5, 9, 6],
   [5, 3, 7, 4]])
    mask = (name=='Javohir') | (name=='Elyor')
    return(data[mask])
    raise NotImplementedError()
```

Your function is checked by the number of elements in the array, the shape of the array `shape`, and whether the elements are numpy arrays. There are also a number of secret tests (our advice is not to try to cheat).

<!-- grade_id: correct_boolean_idx, points: 2, grade: true, solution: false -->
```python

"""小heck the function"""
bool_idx = boolean_slicing() # truncated array
assert bool_idx.size == 16  # number of elements
assert type(bool_idx) == np.ndarray # #numpy array 
```

```python

```