<a href="https://colab.research.google.com/github/Dilshodbek23/Data-Science-and-Artificial-Intelligence/blob/main/task(2_1_1)_arrays_ipynb%22.ipynb" target="_parent"><img src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open In Colab"/></a>

```python
# calling a NumPy library
import numpy as np
```

### Task (1): Create one-dimensional arrays (4 points)
Problem condition: Create a function that creates a one-dimensional array of n elements (n> 0, integers).

Help: use an array or a random function. 


```python
from google.colab import drive
drive.mount('/content/drive')
```

<!-- grade_id: one_dim, points: 0, grade: false, solution: true -->
```python
def one_dim(n):
    """ One-dimensional array of n elements
    """
     # YOUR CODE HERE
    list1 = []
    for i in range(n):
        list1.append(i)
    array1 = np.array(list1)
    return array1 
    raise NotImplementedError()
```

Your function is checked by the number of elements, the size of the ndim array and the numpy array: check your function in the cell below

<!-- grade_id: correct_one_dim, points: 4, grade: true, solution: false -->
```python
"""Check that squares returns the correct output for several inputs"""
for i in range(1, 10):
    assert len(one_dim(i)) == i  # n is the number of elements
    arr = one_dim(i) # massive
    assert arr.ndim == 1 # array size
    assert type(arr) == np.ndarray # #numpy array 
```

### Task (2) Creating two-dimensional arrays (6 points)
Problem: Create a function that creates a 2D array (array) of the shape (r, c).

Help: you can use any of the array creation functions in the numpy library.


<!-- grade_id: two_dim, points: 0, grade: false, solution: true -->
```python
def two_dim(r, c):
    """ two-dimensional array consisting of "r" rows and "c" columns 
    """
     # YOUR CODE HERE
    list_1 = [[]]*r
    for k in range(c):
        list_1[0].append(k)
    array1 = np.array(list_1)
    return array1 
    raise NotImplementedError()
```

Test your function in the cell below

<!-- grade_id: correct_two_dim, points: 6, grade: true, solution: false -->
```python
arr = two_dim(3, 2) # massive
assert arr.shape == (3, 2) # shape
assert arr.ndim == 2 # array size
assert type(arr) == np.ndarray # #numpy array 
```

```python

```